```python
import sys
from pathlib import Path
root = Path().resolve().parent.parent   
if str(root) not in sys.path:
    sys.path.insert(0, str(root))
from utils.grade_submit import submit_grade
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

<img src="data8logo.png" alt="Data 8 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Лабораторна робота 3: Типи даних і масиви

Ласкаво просимо до лабораторної роботи 3!

До цього часу ми використовували Python для обробки чисел і роботи з таблицями. Однак нам потрібно розглянути **типи даних** Python, щоб глибше зрозуміти, як працювати з даними в Python.

У цій лабораторній роботі ви спочатку дізнаєтеся, як представляти та обробляти ще один фундаментальний тип даних: текст. У Python фрагмент тексту називається *рядком* (*string*). Ви також дізнаєтеся, як працювати з *масивами* (*arrays*) даних. Масив наприклад може містити всі числа від 0 до 100 або всі слова з розділу книги. Насамкінець ви створите таблиці та попрактикуєтеся в їхньому аналізі, використовуючи свої знання про операції з таблицями. У [Довіднику Python](https://www.data8.org/fa26/reference/) міститься інформація, яка стане в нагоді під час виконання цієї лабораторної роботи.


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

```python
# Просто запустіть цю комірку
import numpy as np
import math
from datascience import *
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 1. Текст
Програмування стосується не лише чисел. Текст — один із найпоширеніших типів даних, що використовуються в програмах.

У Python текст представлений **рядковим значенням** (**string value**). Слово «рядок» (string) — це програмистський термін, що позначає послідовність символів. Рядок може містити один символ, слово, речення або цілу книгу.

Щоб відрізнити текстові дані від числових даних або власне коду, ми виділяємо рядки, беручи їх у лапки. Допустимі як одинарні лапки (`'`), так і подвійні лапки (`"`), але типи відкриваючих і закриваючих лапок повинні збігатися. Вміст може бути будь-якою послідовністю символів, включаючи цифри та символи.

Ми вже бачили рядки раніше в операторах `print`.  Нижче два різних рядки передаються як аргументи функції `print`. Зверніть увагу, що функція `print` автоматично вставляє пробіл між двома вхідними аргументами-рядками (розділеними комою) у вихідних даних цієї комірки.

```python
print("I <3", 'Data Science')
```

Так само, як числам можна присвоювати імена, імена можна присвоювати й рядковим значенням. Рядки та їхні імена не обов’язково мають бути схожими між собою. Будь-якому рядку можна присвоїти будь-яке ім’я.

```python
one = 'two'
plus = '*'
print(one, plus, one)
```

---
**Питання 1.1.** Юрій Гагарін був першою людиною, яка здійснила політ у космос. Коли він вийшов із капсули після приземлення на Землю, він [як повідомляється](https://en.wikiquote.org/wiki/Yuri_Gagarin) мав таку розмову з жінкою та дівчинкою, які спостерігали за приземленням:

    The woman asked: "Can it be that you have come from outer space?"
    Gagarin replied: "As a matter of fact, I have!"

У комірці нижче міститься незавершений код.  Заповніть пропуски `...` так, щоб вивести цю розмову за допомогою команд `print` *точно* так, як вона наведена вище.

Примітка: назви змінних `woman_asking` та `woman_quote` означають :
- "woman asking" - "запитання жінки"
- "woman quote" - "цитата запитання жінки" (цитати пишуться в лапках, і в цьому випадку лапки є частиною текста)


```python
woman_asking = ...
woman_quote = '"Can it be that you have come from outer space?"'
gagarin_reply = 'Gagarin replied:'
gagarin_quote = ...

print(woman_asking, woman_quote)
print(gagarin_reply, gagarin_quote)
```

```python
grader.check("q11")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.1. Методи роботи з рядками

Strings can be transformed using **methods**. Recall that methods and functions are not the same thing. Here is the textbook section on string methods: [4.2.1 String Methods](https://inferentialthinking.com/chapters/04/2/1/String_Methods.html). 

Here's a sketch of how to call methods on a string:

    <expression that evaluates to a string>.<method name>(<argument>, <argument>, ...)
    
One example of a string method is `replace`, which replaces all instances of some part of the original string (or a *substring*) with a new string. 

    <original string>.replace(<old substring>, <new substring>)
    
`replace` returns (evaluates to) a new string, leaving the original string unchanged.
    
Try to predict the output of this example, then run the cell!

Рядки можна змінювати певним чином за допомогою **методів** (**methods**). Нагадаємо, що методи та функції — це не одне й те саме. Ось розділ підручника, присвячений методам роботи з рядками: [4.2.1 Методи роботи з рядками](https://inferentialthinking.com/chapters/04/2/1/String_Methods.html).

Ось схема виклику методів для рядка:

    <вираз, що має значення у вигляді рядка>.<назва методу>(<аргумент>, <аргумент>, ...)

Одним із прикладів методу для рядків є `replace`, який замінює всі входження певної частини вихідного рядка (цю частину називають *підрядок*, *substring*) на новий рядок.

    <оригінальний рядок>.replace(<старий підрядок>, <новий підрядок>)

`replace` повертає (результатом обчислення є) новий рядок, залишаючи оригінальний рядок без змін.

Спробуйте передбачити результат цього прикладу, а потім запустіть комірку!

```python
# Наступний код замінить в рядку 'bag' одну букву (підрядок) на іншу (новий підрядок)
bag = 'bag'
print(bag.replace('g', 't'), bag)
```

Ви можете викликати функції, застосовуючи їх до результатів інших функцій. Наприклад, `max(abs(-5), abs(3))` дає у результаті 5. Аналогічно, ви можете викликати методи, застосовуючи їх до результатів викликів інших методів або функцій.

Можливо, ви вже помітили одну відмінність між функціями та методами — перед викликом функції, такої як `max`, не потрібно ставити `.`, на відміну від методу для рядків, такого як `replace`. Довідник Python, про який йшлося у вступному абзаці, може стати в нагоді, якщо ви не впевнені, як правильно викликати функцію чи метод.

```python
# Використання методу replace на результаті іншого виклику replace
# Послідовність дій є такою: 
# 1) спочатку в рядку 'train' замінюємо 't' на 'ing', ця дія повертає рядок 'ingrain' 
# 2) в рядку 'ingrain' замінюємо 'in' на 'de'
'train'.replace('t', 'ing').replace('in', 'de')
```

Ось ілюстрація того, як Python виконує «ланцюговий» виклик методу replace:

<img src="chaining_method_calls.png"/>

---
**Питання 1.1.1.** Використовуйте функцію `replace`, щоб перетворити рядок `“hitchhiker”` на `“matchmaker”`. Присвойте отриманий результат змінній `new_word`.

```python
new_word = ...
new_word
```

```python
grader.check("q111")
```

У Python існує ще багато інших методів роботи зі строками, але більшість програмістів не запам’ятовують їхні назви чи способи використання. У «реальному світі» люди зазвичай просто шукають в Інтернеті документацію та приклади. Повний [перелік методів роботи зі строками](https://docs.python.org/3/library/stdtypes.html#string-methods) наведено в документації до мови Python. [Stack Overflow](http://stackoverflow.com) містить величезну базу даних відповідей на запитання, які часто демонструють, як використовувати ці методи для досягнення різних цілей. Матеріал, висвітлений у цих джерелах, який ще не був розглянутий на цьому курсі, не входить до програми курсу.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.2. Конвертація в рядки та із рядків

Рядки та числа — це різні *типи* значень (*types* of values), навіть якщо рядок містить цифри числа. Наприклад, обчислення вмісту наступної комірки призведе до помилки, оскільки ціле число не можна додавати до рядка.

_Примітка:_ Ціле число (an integer) не можна додавати до рядка, це немає сенсу, так само як і число з плаваючою комою (a float). Однак цілі числа та числа з плаваючою комою **можна** додавати між собою!

```python
8 + "8"
```

Однак існують вбудовані функції для перетворення чисел у рядки та рядків у числа. Деякі з цих вбудованих функцій мають обмеження щодо типу аргументів, які вони приймають:

|Функція |Опис|
|-|-|
|`int`|Перетворює рядок цифр або число з плаваючою комою на ціле значення («int»)|
|`float`|Перетворює рядок цифр (можливо з десятковою крапкою) або ціле число на число з плаваючою комою («float»)|
|`str`|Перетворює будь-яке значення на рядок|

Спробуйте передбачити, до якого типу даних і якого значення обчислюється `example`, а потім запустіть комірку.

```python
example = 8 + int("10") + float("8")

print(example)
print("This example returned a " + str(type(example)) + "!")
```

Припустимо, ви пишете програму, яка шукає дати в тексті, і хочете, щоб вона обчислювала проміжок часу, що минув між двома виявленими роками.  Віднімати два тексти не має сенсу, але ви можете спочатку перетворити текст, що містить роки, на числа.

---

**Питання 1.2.1.** Допишіть наведений нижче код, щоб обчислити кількість років, що минула між `one_year` та `another_year`.  Не пишіть просто числа `1618` та `1648` (або `30`); скористайтеся функцією перетворення, щоб перетворити задані текстові дані на числа.

```python
# Деякі текстові дані:
one_year = "1618"
another_year = "1648"

# Заповніть наступний рядок.  Зверніть увагу, що ми не можемо просто написати:
# another_year - one_year
# Якщо ви не розумієте чому, спробуйте побачити, що станеться, коли ви
# напишіть це тут.
difference = ...
difference
```

```python
grader.check("q121")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1.3. Передача (passing) рядків у функції

Значення рядків, як і числа, можуть бути аргументами функцій і можуть повертатися функціями як результат.

Функція `len` (походить від слова «length», «довжина») приймає один рядок як аргумент і повертає кількість символів (включно з пробілами та розділовими знаками) у цьому рядку.

Зверніть увагу, що вона не підраховує кількість *слів*. Вираз `len("one small step for man")` обчислюється як 22 символи, а не 5 слів.

**Запитання 1.3.1.**  Використовуйте функцію `len`, щоб визначити кількість символів у довгому рядку в наступній комірці.  До символів належать, зокрема, пробіли та розділові знаки. Результат (кількість символів) збережіть в змінній з ім'ям `sentence_length`.

(Цей рядок — це перше речення англійського перекладу французької [Декларації прав людини](http://avalon.law.yale.edu/18th_century/rightsof.asp).)


```python
a_very_long_sentence = "The representatives of the French people, organized as a National Assembly, believing that the ignorance, neglect, or contempt of the rights of man are the sole cause of public calamities and of the corruption of governments, have determined to set forth in a solemn declaration the natural, unalienable, and sacred rights of man, in order that this declaration, being constantly before all the members of the Social body, shall remind them continually of their rights and duties; in order that the acts of the legislative power, as well as those of the executive power, may be compared at any moment with the objects and purposes of all political institutions and may thus be more respected, and, lastly, in order that the grievances of the citizens, based hereafter upon simple and incontestable principles, shall tend to the maintenance of the constitution and redound to the happiness of all."
sentence_length = ...
sentence_length
```

```python
grader.check("q131")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# 2. Arrays

Computers are most useful when you can use a small amount of code to *do the same action* to *many different things*.

For example, in the time it takes you to calculate the 18% tip on a restaurant bill, a laptop can calculate 18% tips for every restaurant bill paid by every human on Earth that day.  (That's if you're pretty fast at doing arithmetic in your head!)

**Arrays** are how we put many values in one place so that we can operate on them as a group. For example, if `billions_of_numbers` is an array of numbers, the expression

    .18 * billions_of_numbers

gives a new array of numbers that contains the result of multiplying **each number** in `billions_of_numbers` by .18.  Arrays are not limited to numbers; we can also put all the words in a book into an array of strings.

Concretely, an array is a **collection of values of the same type**. 

# 2. Масиви (Arrays)

Комп’ютери найкорисніші тоді, коли за допомогою невеликої кількості коду можна *виконати одну й ту саму дію* над *багатьма різними об’єктами*.

Наприклад, за той час, що вам знадобиться, щоб обчислити в умі 18% чайових до рахунку в ресторані, ноутбук може обчислити 18% чайових для кожного рахунку в ресторані, оплаченого кожною людиною на Землі цього дня. (Звісно, якщо ви досить швидко рахуєте в умі!)

**Масиви** (**Arrays**) — це спосіб розміщення багатьох значень в одному місці, щоб ми могли оперувати ними як групою. Наприклад, якщо `billions_of_numbers` — це масив чисел, то вираз

    .18 * billions_of_numbers

дає новий масив чисел, що містить результат множення **кожного числа** у `billions_of_numbers` на .18.  Масиви не обмежуються числами; ми також можемо помістити всі слова з книги в масив рядків.

Конкретно кажучи, масив — це **колекція значень одного типу**.

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" /> 

## 2.1. Створення масивів

Спочатку давайте дізнаємося, як вручну вводити значення в масив. Зазвичай в програмах так не роблять. Як правило, ми створюємо масиви, завантажуючи їх із зовнішнього джерела, наприклад, із файлу даних.

Щоб створити масив вручну, викличте функцію `make_array`. Кожен аргумент, який ви передаєте функції `make_array`, потрапить до масиву, який вона повертає. Виконайте цю комірку, щоб побачити приклад:

```python
make_array(0.125, 4.75, -1.3)
```

Кожне значення в масиві (у наведеному вище прикладі — числа 0,125, 4,75 та -1,3) називається *елементом* (*element*) цього масиву.

Самі масиви також є значеннями, так само як числа та рядки.  Це означає, що їм можна присвоювати (assign) іменам або їх можна використовувати як аргументи функцій. Наприклад, `len(<some_array>)` повертає кількість елементів у `some_array`.

**Завдання 2.1.1.** Створіть масив, що містить числа 0, 1, -1 та $\pi$ у наведеному порядку.  Назвіть його `interesting_numbers`.

*Підказка:* Згадайте як ви отримали значення $\pi$ у лабораторній роботі 2?  Тут ви можете скористатися ним таким самим чином. Модуль `math` було імпортовано на початку цього ноутбука, його імпортувати не треба.


```python
interesting_numbers = ...
interesting_numbers
```

```python
grader.check("q211")
```

--- 
**Question 2.1.2.** Make an array containing the five strings `"Hello"`, `","`, `" "`, `"world"`, and `"!"`.  (The third one is a single space inside quotes.)  Name it `hello_world_components`.

*Note:* If you evaluate `hello_world_components`, you'll notice some extra information in addition to its contents: `dtype='<U5'`.  That's just NumPy's extremely cryptic way of saying that the data types in the array are strings.

---
**Завдання 2.1.2.** Створіть масив, що містить п’ять символьних рядків: `"Hello"`, `","`, `" "`, `"world"`, and `"!"`.  (Третій рядок — це один пробіл у лапках.)  Назвіть його `hello_world_components`.

*Примітка:* Якщо ви обчислите `hello_world_components`, то, крім його вмісту, ви помітите додаткову інформацію: `dtype=“<U5”`.  Це просто надзвичайно загадковий спосіб, яким бібліотека NumPy повідомляє, що типи даних у масиві — це рядки.



```python
hello_world_components = ...
hello_world_components
```

```python
grader.check("q212")
```

###  `np.arange`
Масиви в Python забезпечуються пакетом під назвою [NumPy](http://www.numpy.org/) (вимовляється як «NUM-pie»). Пакет називається `numpy`, але зазвичай його перейменовують на `np` для стислості.  Це можна зробити так:

    import numpy as np

У науці про дані дуже часто доводиться працювати з великою кількістю чисел, рівномірно розподілених у певному діапазоні.  NumPy надає для цього спеціальну функцію під назвою `arange`.  Рядок коду `np.arange(start, stop, step)` створює масив, що містить усі числа, починаючи з `start` і збільшуючись на `step`, зупиняючись **до** досягнення `stop`.

Виконайте наведені нижче комірки, щоб побачити кілька прикладів!

```python
# Цей масив починається з 1 (start) і збільшується на 2 (step)
# а потім зупиняється перед 6 (stop)
np.arange(1, 6, 2)
```

```python
# Цей масив не містить числа 9,
# оскільки функція np.arange зупиняється *до* досягнення значення stop
np.arange(4, 9, 1)
```

```python
# Значенням за замовчуванням для параметра start у функції np.arange є 0
# Значенням за замовчуванням для параметра step у функції np.arange є 1
# Цей масив за замовчуванням починається з 0 і збільшується на 1, але зупиняється до досягнення 5,
# оскільки функції np.arange було передано один аргумент, 
# для параметрів start і step використовуються наведені вище значення за замовчуванням.
np.arange(5)
```

---
**Питання 2.1.3.** Імпортуйте `numpy` як `np`, а потім за допомогою `np.arange` створіть масив, що містить числа кратні 99 (the multiples of 99), від 0 до (**включно**) 9999.  (Тобто його елементами будуть 0, 99, 198, 297 тощо, 9999.)

```python
...
multiples_of_99 = ...
multiples_of_99
```

```python
grader.check("q213")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.2. Робота з окремими елементами масивів («індексація», "indexing")
Давайте попрацюємо з більш цікавим набором даних. У наступній комірці створюється масив під назвою `population_amounts`, який містить оціночні дані щодо чисельності населення світу за кожен рік, починаючи з **1950** року і приблизно до сьогодення. (Ці оціночні дані взято з веб-сайту Бюро перепису населення США.)

Замість того, щоб вводити дані вручну, ми завантажили їх із файлу на вашому комп’ютері під назвою `world_population.csv`.  Ви дізнаєтеся, як це зробити, пізніше в цій лабораторній роботі!

```python
population_amounts = Table.read_table("world_population.csv").column("Population")
population_amounts
```

Ось як ми отримуємо перший елемент масиву `population_amounts`, який відображає чисельність населення світу у першому році набору даних — 1950 році.

Примітка: "item" можна перекласти як "елемент"

```python
population_amounts.item(0)
```

The value of that expression is the number 2,557,628,654 (around 2.5 billion), because that's the first thing in the array `population_amounts`.

Notice that we wrote `.item(0)`, not `.item(1)`, to get the first element.  This is a weird convention in computer science.  **0 is called the *index* of the first item.**  It's the number of elements that appear *before* that item.  So 3 is the index of the 4th item.

Here are some more examples.  In the examples, we've given names to the things we get out of `population_amounts`.  Read and run each cell.
Значенням цього виразу є число 2,557,628,654 (приблизно 2,5 мільярда), оскільки це перший елемент масиву `population_amounts`.

Зверніть увагу, що для отримання першого елемента ми написали `.item(0)`, а не `.item(1)`. Це дивна традиція в інформатиці.  **0 називається *індексом* першого елемента.**  Цей індекс елемента показує кількість елементів, що знаходяться *перед* цим елементом.  Отже, 3 — це індекс четвертого елемента (перед ним знаходяться 3 елемнети).

Ось ще кілька прикладів.  У прикладах ми дали імена елементам, які отримуємо з `population_amounts`.  Прочитайте та запустіть кожну комірку.

```python
# Чисельність населення в 1962 році є 13-м елементом масиву
# тому що це 12 індексів після 1950 року, тобто індекс 0. (1950 + 12 = 1962)

population_1962 = population_amounts.item(12)
population_1962
```

```python
# 66-й елемент – чисельність населення у 2015 році.
population_2015 = population_amounts.item(65)
population_2015
```

```python
# Масив має лише 66 елементів, тому це не працює.
# (Немає елемента з 66 іншими елементами перед ним.)
population_2016 = population_amounts.item(66)
population_2016
```

Оскільки `make_array` повертає масив, ми можемо викликати метод `.item(3)` на його результаті, щоб отримати його 4-й елемент, так само, як раніше ми з'єднали ланцюжком виклики методу `replace`.

```python
make_array(-1, -3, 4, -2).item(3)
```

---
**Питання 2.2.1.** Присвоїти змінній `population_1973` значення, що відповідає чисельності населення світу у 1973 році, отримавши відповідний елемент із масиву `population_amounts` за допомогою функції `item`.

```python
population_1973 = ...
population_1973
```

```python
grader.check("q221")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 2.3. Виконання дії над кожним елементом масиву
Масиви в першу чергу корисні для багатократного виконання однієї й тієї ж операції над усіма елементами, завдяки цьому нам не доводиться часто використовувати `.item` і працювати з окремими елементами.

### Округлення (Rounding)
Ось одне просте запитання, яке ми можемо поставити щодо чисельності населення світу:

> Якою була чисельність населення кожного року, округлена до найближчого мільйона?

Округлення часто використовується з великими числами, коли нам не потрібна надто висока точність у цифрах. Одним із прикладів цього є представлення даних у таблицях та візуалізаціях.

Ми можемо спробувати відповісти на наше запитання, використовуючи вбудовану в Python функцію `round` та метод `item`, який ви щойно побачили.

**Примітка:** функція `round` приймає два аргументи: число, яке потрібно округлити, та кількість знаків після коми, до якої слід округлити. Другий аргумент можна розглядати як кількість кроків вправо або вліво від десяткової коми. Від’ємні числа означають переміщення вліво, а додатні — вправо. Отже, якщо ми маємо `round(1234.5, -2)`, це означає, що слід переміститися на два знаки вліво, а потім замінити всі цифри праворуч від цього знака на нулі. У результаті отримаємо число 1200.0. Такий самий результат ми отримаємо і в цьому випадку `round(1234, -2)` (кома хоч вона тут і не написана явно, розташована після цифри `4`, тобто `1234.0`) З іншого боку, якщо ми маємо `round(6.789, 1)`, нам слід переміститися на одну позицію вправо, а потім замінити всі цифри праворуч від цієї позиції на нулі. У результаті буде виведено число 6.8.


<img src="np_round_example.png">

```python
population_1950_magnitude = round(population_amounts.item(0), -6)
population_1951_magnitude = round(population_amounts.item(1), -6)
population_1952_magnitude = round(population_amounts.item(2), -6)
population_1953_magnitude = round(population_amounts.item(3), -6)
print(population_1950_magnitude, population_1951_magnitude, population_1952_magnitude, population_1953_magnitude)

```

Але це нудно і не дозволяє повною мірою скористатися тим, що ми працюємо з комп’ютером.

Натомість NumPy надає власну версію функції `round`, яка округлює кожен елемент масиву. Вона приймає два аргументи: один - масив чисел та другий - кількість знаків після коми, до якої слід округлювати. Вона повертає масив тієї самої довжини, де перший елемент результату — це округлений перший елемент аргументу, і так далі.

---

**Завдання 2.3.1.** Використовуйте `np.round`, щоб обчислити чисельність населення світу для кожного року заокруглену до найближчого мільйона (6 нулів). Назвіть результат (масив із 66 чисел) `population_rounded`.  Ваш код має бути дуже коротким.

```python
population_rounded = ...
population_rounded
```

```python
grader.check("q231")
```

Те, що ви щойно зробили, називається **поелементним** (**elementwise**) застосуванням функції `np.round`, оскільки `np.round` діє окремо на кожен елемент масиву, до якого вона застосовується. Ось ілюстрація того, що відбувається:

<img src="array_round.jpg">

У [розділі](https://www.inferentialthinking.com/chapters/05/1/Arrays) підручника про масиви є корисний список функцій NumPy, розроблених для поелементної роботи, наприклад `np.round`.

### Арифметика
Арифметичні операції також виконуються поелементно, тобто якщо ви виконуєте арифметичну операцію (наприклад, віднімання, ділення тощо) над масивом, Python виконає цю операцію окремо для кожного елемента масиву та поверне масив, що містить усі результати. Наприклад, ви можете поділити всі значення чисельності населення на 1 мільярд, щоб отримати числа в мільярдах:

```python
population_in_billions = population_amounts / 1000000000
population_in_billions
```

Ви можете зробити те саме з додаванням, відніманням, множенням і піднесенням до степеня (`**`).

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 3. Створення таблиць

Масив зручно використовувати для опису одного параметра кожного елемента в колекції. Наприклад, припустимо, що наша колекція складається з усіх штатів США. Тоді масив міг би описувати площу кожного штату (один з можливих параметрів).

Таблиці розширюють цю ідею, містячи кілька стовпців (columns), що зберігаються та представлені у вигляді масивів, кожен з яких описує окремий атрибут (параметр) для кожного елемента колекції. Таким чином, таблиці дозволяють нам не тільки зберігати дані про багато об’єктів, але й містити кілька видів даних про кожен об’єкт.

Наприклад, у комірці нижче ми маємо два масиви. Перший, `population_amounts`, був визначений вище в розділі 2.2 і містить дані про чисельність населення світу за кожним роком (за оцінками Бюро перепису населення США). Другий масив, `years`, містить самі роки. Ці елементи впорядковані, тому рік і чисельність населення світу за цей рік мають однаковий індекс у відповідних масивах.

```python
# Просто запустіть цю клітинку

years = np.arange(1950, 2015+1)
print("Стовпець чисельності населення:", population_amounts)
print("Стовпець років:", years)
```

Припустимо, ми хочемо відповісти на таке запитання:

> У якому році чисельність населення світу перевищила 6 мільярдів?

Технічно на це запитання можна відповісти, просто вручну аналізуючи данні в масивах (див. результати виведені на екран після останньої комірки з кодом), але це дещо заплутано, оскільки доведеться підрахувати позицію, на якій чисельність населення вперше перевищила 6 мільярдів, а потім знайти відповідний елемент у масиві років. У таких випадках простіше буде помістити дані в *`Table`* — двовимірний тип набору даних.

Наведений нижче вираз:

- створює порожню таблицю за допомогою виразу `Table()`,
- додає два стовпці, викликаючи `with_columns` із чотирма аргументами (що відповідають назві стовпця, значенням для цього стовпця, назві другого стовпця та значенням для цього стовпця),
- присвоює результат імені `population`, і, нарешті,
- обчислює `population`, щоб ми могли побачити таблицю.

Строки `«Year»` та `«Population»` — це назви стовпців, які ми обрали. Імена `population_amounts` та `years` були присвоєні вище двом масивам **однакової довжини**. Функція `with_columns` (документацію можна знайти [тут](http://data8.org/fa26/reference)) приймає по черзі такі пари аргументів: рядок (що представляє назву стовпця) та масив (що представляє дані в цьому стовчику). Рядки та масиви розділяються комами. Таких пар має бути стільки скільки в таблиці буде стовпчиків.

_Примітка:_ Якщо ви хочете додати один стовпець до таблиці, ви можете використовувати або `with_columns`, або `with_column` з двома аргументами (назва стовпця та масив). Однак, якщо ви хочете додати до таблиці більше одного стовпця, ви повинні використовувати `with_columns` з чергуванням аргументів у вигляді рядків та масивів.

```python
population = Table().with_columns(
    "Population", population_amounts,
    "Year", years
)
population
```

Тепер дані об’єднано в одну таблицю! Аналізувати ці дані стало набагато простіше. Наприклад, якщо вам потрібно дізнатися, якою була чисельність населення у 1959 році, ви зможете це побачити з першого погляду.

---
**Питання 3.1.** У комірці нижче ми створили 2 масиви. Використовуючи наведені вище кроки, присвойте імені `top_10_movies` таблицю, що має два стовпці з назвами "Name" та "Rating", які містять відповідно `top_10_movie_names` та `top_10_movie_ratings`.


```python
top_10_movie_names = make_array(
        'The Shawshank Redemption (1994)',
        'The Godfather (1972)',
        'The Godfather: Part II (1974)',
        'Pulp Fiction (1994)',
        "Schindler's List (1993)",
        'The Lord of the Rings: The Return of the King (2003)',
        '12 Angry Men (1957)',
        'The Dark Knight (2008)',
        'Il buono, il brutto, il cattivo (1966)',
        'The Lord of the Rings: The Fellowship of the Ring (2001)')
top_10_movie_ratings = make_array(9.2, 9.2, 9., 8.9, 8.9, 8.9, 8.9, 8.9, 8.9, 8.8)

top_10_movies = ...

# Ми розмістили тут наступний рядок 
# щоб ваша таблиця була виведена на екран
# під час запуску цієї комірки.
top_10_movies
```

```python
grader.check("q31")
```

#### Завантаження таблиці з файлу

У більшості випадків ми не будемо витрачати час на введення всіх даних вручну. Натомість ми завантажуємо їх із зовнішнього джерела, наприклад, із файлу даних. Існує багато форматів файлів даних, але найпоширенішим є CSV ("comma-separated values", "значення, розділені комами").

`Table.read_table(...)` приймає один аргумент (шлях до файлу даних у форматі **рядка**) і повертає таблицю.

**Завдання 3.2.** Файл `imdb.csv` містить таблицю з інформацією про 250 фільмів із найвищими рейтингами на IMDb.  Завантажте її як таблицю з назвою `imdb`.

(Можливо, ви пам’ятаєте, як працювали з цією таблицею в лабораторній роботі 2!)


```python
imdb = ...
imdb
```

```python
grader.check("q32")
```

Звідки взявся файл `imdb.csv`? Загляньте в [папку цього лабораторного завдання](./). Ви повинні побачити файл із назвою `imdb.csv`.

Відкрийте файл `imdb.csv` у цій папці та ознайомтеся з його форматом. Що ви помітили? Розширення імені файлу `.csv` вказує на те, що цей файл має [формат CSV (значення, розділені комами)](http://edoceo.com/utilitas/csv-file-format).

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 4. Ще більше операцій з таблицями!

Тепер, коли ви вже попрацювали з масивами, давайте додамо ще кілька методів до переліку операцій з таблицями, з якими ви ознайомилися в лабораторній роботі 2.

### `tbl.column`

`tbl.column` приймає як аргумент ім’я стовпця (the column name) таблиці (у вигляді рядка) і повертає значення цього стовпця у вигляді **масиву**.

```python
# Повертає масив назв фільмів
top_10_movies.column('Name')
```

### `take`
Метод таблиці `take` приймає як аргумент масив чисел. Кожне число має бути індексом рядка в таблиці. Він повертає **нову таблицю**, що містить лише ці рядки.

Зазвичай метод `take` використовують разом із `np.arange`, щоб вибрати перші кілька рядків таблиці.

```python
# Зніміть перші 5 фільмів із top_10_movies
top_10_movies.take(np.arange(0, 5, 1))
```

Наступні три завдання допоможуть вам попрактикуватися в поєднанні операцій, які ви вивчили в цій та попередній лабораторних роботах, для відповіді на запитання щодо таблиць `population` та `imdb`. Спочатку ознайомтеся з таблицею `population` із розділу 2.

```python
# Запустіть цю комірку, щоб відобразити таблицю населення.
population
```

---
**Питання 4.1.** Ознайомтеся з таблицею `population` із розділу 2 цієї лабораторної роботи.  Обчисліть рік, коли чисельність населення світу вперше перевищила 6 мільярдів. Запишіть цей рік у змінну `year_population_crossed_6_billion`.

_Підказка:_ 1) Як можна відфільтрувати таблицю, щоб залишити лише рядки зі значеннями населення понад 6 мільярдів? 2) Як знайти перший рік у цій відфільтрованій таблиці? Ознайомтеся з деякими методами роботи з таблицями в [Довідковому листі Python](https://www.data8.org/fa26/reference/)!

```python
year_population_crossed_6_billion = ...
year_population_crossed_6_billion
```

```python
grader.check("q41")
```

---
**Завдання 4.2.** Визначте середню оцінку фільмів, що вийшли до 2000 року, та середню оцінку фільмів, що вийшли у 2000 році або пізніше, для фільмів у базі даних `imdb`.

*Підказка*: Подумайте, які кроки вам потрібно виконати (обчислити середнє значення, знайти оцінки, знайти фільми, випущені у 20-му/21-му століттях), і спробуйте розташувати їх у логічній послідовності.


```python
before_2000 = ...
after_or_in_2000 = ...
before_2000_avg = ...
after_or_in_2000_avg = ...
print("Average before 2000 rating:", before_2000_avg)
print("Average after or in 2000 rating:", after_or_in_2000_avg)
```

```python
grader.check("q42")
```

---
**Завдання 4.3.** Ось вам завдання: знайдіть кількість фільмів, що вийшли в *парні* роки.

*Підказка:* Оператор `%` обчислює остачу від ділення на число. Отже, `5 % 2` дорівнює 1, а `6 % 2` — 0. Число є парним, якщо при діленні на 2 остача дорівнює 0.

*Підказка 2:* Оператор `%` можна застосовувати до масивів, і він діє по елементах, як `+` чи `*`.  Отже, `make_array(5, 6, 7) % 2` дорівнює `array([1, 0, 1])`.

*Підказка 3:* Додайте новий стовпець під назвою «Year Remainder», який міститиме залишок від ділення року виходу кожного фільму на 2. Для цього можна використати `tbl.with_column(col_name, col_values)`: метод таблиці, який приймає ім’я нового стовпця та масив значень і повертає копію вихідної таблиці `tbl` із новим стовпцем.  Потім скористайтеся `where`, щоб знайти рядки, у яких значення цього нового стовпця дорівнює 0.  Нарешті, скористайтеся `num_rows`, щоб підрахувати кількість таких рядків.

*Примітка:* Ці кроки можна ланцюжком об’єднати в одну інструкцію (statement) або розбити на кілька рядків, присвоївши проміжні імена. Ви завжди можете розбивати задачі так, як вам зручно!

```python
num_even_year_movies = ...
num_even_year_movies
```

```python
grader.check("q43")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 5. Додаткові вправи з масивами

Наведені нижче завдання стануть чудовою нагодою попрактикуватися у створенні масивів різними способами, а також у застосуванні різних методів арифметичних операцій з масивами!


###  5.1 `np.arange` (продовження)

##### Дані про температуру
NOAA (Національне управління з океанів і атмосфери США) експлуатує метеостанції, які вимірюють температуру поверхні в різних точках Сполучених Штатів. Якщо вам цікаво, щогодинні дані [доступні для загального огляду](http://www.ncdc.noaa.gov/qclcd/QCLCD?prior=N).

Припустимо, що ми завантажуємо всі погодинні дані зі станції в Окленді, штат Каліфорнія, за грудень 2015 року.  Для аналізу даних нам потрібно знати, коли було зроблено кожне вимірювання, але виявляється, що дані не містять часових міток вимірювань (часу, коли кожне з них було зроблено).

Однак ми знаємо, що перше вимірювання було зроблено в перший момент грудня 2015 року (опівночі 1 грудня), а кожне наступне — рівно через 1 годину після попереднього.

---
**Завдання 5.1.1.** Створіть масив, що містить *час у секундах від початку місяця*, у який було зроблено **кожне погодинне вимірювання**.  Назвіть його `collection_times`.

*Підказка 1:* У грудні було 31 день, що еквівалентно ($31 \times 24$) годинам або ($31 \times 24 \times 60 \times 60$) секундам.  Отже, ваш масив повинен містити $31 \times 24$ елементів.

*Підказка 2:* Функція `len` також працює з масивами!  Якщо ваш масив `collection_times` не проходить тести, перевірте його довжину та переконайтеся, що він містить $31 \times 24$ елементів.

```python
collection_times = ...
collection_times
```

```python
grader.check("q511")
```

### 5.2 Виконання операції над усіма елементами масиву (продовження)

Обчисліть суму чайових за кількома рахунками в ресторані одночасно (у цьому випадку — лише за 3ма):

```python
restaurant_bills = make_array(20.12, 39.90, 31.01)
print("Ресторанні рахунки:\t", restaurant_bills)

# Множення масиву
tips = .2 * restaurant_bills
print("Чайові:\t\t\t", tips)
```

<img src="array_multiplication.jpg">

---

**Питання 5.2.1** Припустимо, що загальна сума рахунку в ресторані дорівнює сумі початкового рахунку та чайових. Якщо чайові становлять 20 %, це означає, що для отримання загальної суми рахунку потрібно помножити початковий рахунок на 1.2.  Обчисліть загальну суму рахунку для кожного рахунку в масиві `restaurant_bills` і присвойте отриманий масив змінній `total_charges`.


```python
total_charges = ...
total_charges
```

```python
grader.check("q521")
```

---

**Питання 5.2.2.** Масив `more_restaurant_bills` містить 100 000 рахунків!  Обчисліть загальну суму кожного рахунку в масиві `more_restaurant_bills`.  Чим відрізняється ваш код?


```python
more_restaurant_bills = Table.read_table("more_restaurant_bills.csv").column("Bill")
more_total_charges = ...
more_total_charges
```

```python
grader.check("q522")
```

Функція `sum` приймає в якості аргументу один масив чисел. Вона повертає суму всіх чисел у цьому масиві (тобто повертає одне число, а не масив).

---

**Питання 5.2.3.** Якою була сума всіх рахунків у масиві `more_restaurant_bills`, *включно з чайовими*?


```python
sum_of_bills = ...
sum_of_bills
```

```python
grader.check("q523")
```

---

**Питання 5.2.4.** Степені числа 2 ($2^0 = 1$, $2^1 = 2$, $2^2 = 4$ тощо) часто зустрічаються в інформатиці.  (Наприклад, ви, мабуть, помічали, що обсяг пам’яті на смартфонах або USB-накопичувачах вимірюється в степенях числа 2, таких як 16 ГБ, 32 ГБ або 64 ГБ.)  Використовуйте `np.arange` та оператор піднесення до степеня `**`, щоб обчислити перші 30 степенів числа 2, починаючи з `2^0`.

*Підказка 1:* `np.arange(1, 2**30, 1)` створює масив із $2^{30}$ елементами і **призведе до збою обчислювального ядра (kernel)** - не робіть так.

*Підказка 2:* Частина вашого рішення буде пов’язана з використанням `np.arange`, але ваш масив не повинен містити більше 30 елементів.

```python
powers_of_2 = ...
powers_of_2
```

```python
grader.check("q524")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

```python
secret_jupytutor = ...
```

## «Улюбленці даних 8»

**Преслі** дуже пишається тобою за те, що ти завершив лабораторну роботу № 03!

<img src="presley.jpg" width="40%" />



---

Ви завершили лабораторну роботу!

**Важлива інформація щодо подання роботи:**
- **Виконайте всі тести** і переконайтеся, що всі вони виконані успішно
- **Збережіть** роботу через меню **«Файл»**

**Ви несете відповідальність за те, щоб переконатися, що ваша робота збережена, перш ніж запускати останню комірку.**


---

Щоб ще раз перевірити свою роботу, клітинка нижче повторно запустить усі тести автогрейдера.

```python
grader.check_all()
```

## Подання

Перед виконанням наведеної нижче комірки переконайтеся, що ви виконали всі комірки у своєму зошиті послідовно, щоб у результаті відобразилися всі зображення та графіки. Після виконання двох наведених нижче коміркок ваші результати буде передано викладачу.

#### 1) Напишіть своє ім'я та призвище
В наступній комірці впишіть свої власні ім'я та призвище.  
Запустіть комірку

```python
# змінній student_name присвойте значення вашого імені та прізвища, як у прикладі нижче.
# student_name = "Андрій Ковальчук"
student_name = 
```

#### 2) Наступну комірку просто запустіть

```python
lab_name = "lab03"
comment = "arrays"
r = grader.check_all()
submit_grade(student_name, lab_name, comment, grader.check_all(), True)
```

Tests for q11 (points: None)
```
>>> woman_asking
'The woman asked:'
```
```
>>> gagarin_quote
'"As a matter of fact, I have!"'
```


Tests for q111 (points: None)
```
>>> new_word
'matchmaker'
```


Tests for q121 (points: None)
```
>>> abs(difference) == 30 or abs(difference) == 30.0
True
```


Tests for q131 (points: None)
```
>>> sentence_length
896
```


Tests for q211 (points: None)
```
>>> import numpy as np
>>> type(interesting_numbers) == np.ndarray
True
```
```
>>> len(interesting_numbers)
4
```
```
>>> import numpy as np
>>> all(interesting_numbers == np.array([0, 1, -1, math.pi]))
True
```


Tests for q212 (points: None)
```
>>> import numpy as np
>>> type(hello_world_components) == np.ndarray
True
```
```
>>> len(hello_world_components)
5
```
```
>>> import numpy as np
>>> all(hello_world_components == np.array(['Hello', ',', ' ', 'world', '!']))
True
```


Tests for q213 (points: None)
```
>>> type(multiples_of_99) == np.ndarray
True
```
```
>>> len(multiples_of_99)
102
```
```
>>> all(multiples_of_99 == np.arange(0, 9999 + 99, 99))
True
```


Tests for q221 (points: None)
```
>>> population_1973
3942096442
```


Tests for q231 (points: None)
```
>>> import numpy as np
>>> type(population_rounded) == np.ndarray
True
```
```
>>> import numpy as np
>>> sum(population_rounded) == 312868000000
True
```


Tests for q31 (points: None)
```
>>> type(top_10_movies) == tables.Table
True
```
```
>>> top_10_movies.select('Name', 'Rating')
Name                                                     | Rating
The Shawshank Redemption (1994)                          | 9.2
The Godfather (1972)                                     | 9.2
The Godfather: Part II (1974)                            | 9
Pulp Fiction (1994)                                      | 8.9
Schindler's List (1993)                                  | 8.9
The Lord of the Rings: The Return of the King (2003)     | 8.9
12 Angry Men (1957)                                      | 8.9
The Dark Knight (2008)                                   | 8.9
Il buono, il brutto, il cattivo (1966)                   | 8.9
The Lord of the Rings: The Fellowship of the Ring (2001) | 8.8
```


Tests for q32 (points: None)
```
>>> type(imdb) == tables.Table
True
```
```
>>> imdb.num_rows == 250
True
```
```
>>> imdb.select('Votes', 'Rating', 'Title', 'Year', 'Decade').sort(0).take(range(2, 5))
Votes | Rating | Title                       | Year | Decade
31003 | 8.1    | Le salaire de la peur       | 1953 | 1950
32385 | 8      | La battaglia di Algeri      | 1966 | 1960
35983 | 8.1    | The Best Years of Our Lives | 1946 | 1940
```


Tests for q41 (points: None)
```
>>> import numpy as np
>>> type(year_population_crossed_6_billion) == int or type(year_population_crossed_6_billion) == np.int64
True
```
```
>>> year_population_crossed_6_billion == 1999
True
```


Tests for q42 (points: None)
```
>>> abs(before_2000_avg - 8.278362573099415) < 1e-05
True
```
```
>>> abs(after_or_in_2000_avg - 8.237974683544303) < 1e-05
True
```


Tests for q43 (points: None)
```
>>> num_even_year_movies == 127
True
```


Tests for q511 (points: None)
```
>>> type(collection_times) == np.ndarray
True
```
```
>>> len(collection_times)
744
```
```
>>> collection_times.item(12) == 43200
True
```
```
>>> collection_times.item(37) == 133200
True
```


Tests for q521 (points: None)
```
>>> import numpy as np
>>> type(total_charges) == np.ndarray
True
```
```
>>> import numpy as np
>>> sum(abs(total_charges - np.array([24.144, 47.88, 37.212]))) < 1e-06
True
```


Tests for q522 (points: None)
```
>>> import numpy as np
>>> type(more_total_charges) == np.ndarray
True
```
```
>>> import numpy as np
>>> sum(abs(more_total_charges - 1.2 * more_restaurant_bills)) < 1e-06
True
```


Tests for q523 (points: None)
```
>>> round(sum_of_bills, 2) == 1795730.06
True
```


Tests for q524 (points: None)
```
>>> all(powers_of_2 == make_array(1, 2, 4, 8, 16, 32, 64, 128, 256, 512, 1024, 2048, 4096, 8192, 16384, 32768, 65536, 131072, 262144, 524288, 1048576, 2097152, 4194304, 8388608, 16777216, 33554432, 67108864, 134217728, 268435456, 536870912))
True
```
